Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `???` or "YOUR ANSWER HERE".

---

# CSCI 3155 Assignment 3 : Inductive Definitions and Case Pattern Matching.

This assignment asks you to write scala programs. 

**Restrictions** apply to each problem in terms of forbidden Scala features and API functions. Please read them carefully and ask for clarifications from the course staff over Piazza or during office hours if unsure.

Note: `???` indicates that there is a missing function or code fragment that needs to be filled in. In scala, 
it is also a macro that throws a `NotImplemented` exception. Make sure that you remove the `???` and replace it with the answer. 

Use the test cases provided to test them. You are also encouraged to write your own test cases to help debug your work. However, please delete any extra cells you may have created lest they break our autograder.

**Very Important:** Please run the cell that defines the functions `passed` and `testWithMessage` below whenever you restart the notebook.

### Connor O'Reilly

```scala
// TEST HELPER

// FIRST RUN THIS CELL EVERY TIME YOU START THE NOTEBOOK
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

def testWithMessage[T](v1: T, expected: T, testID: String) = { 
    println(s"Test $testID"); 
    println(s"\t Expected: $expected, your code returned: $v1")
    assert (v1 == expected, s"Test $testID FAILED.")
    println("\t Passed!")
}

def testWithMessage(v1: Double, expected: Double, testID: String) = {
    val tolerance = 1E-5
    println(s"Test $testID -- comparing with tolerance $tolerance."); 
    println(s"\t Expected: $expected, your code returned: $v1")
    assert (math.abs(v1-expected) <= tolerance, s"Test $testID FAILED.")
    println("\t Passed!")
}
    
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
defined [32mfunction[39m [36mtestWithMessage[39m
defined [32mfunction[39m [36mtestWithMessage[39m
```

## Problem 1

This problem concerns inductive definitions of _continued fractions_. 
A continued fraction is a number defined as : 

$$ g = \frac{1}{a_0 + \frac{1}{a_1 + \frac{1}{a_2 + \frac{1}{a_3 + \frac{1}{\ddots + \frac{1}{a_n}}}}}}$$ 

### Example

Consider this fraction:

$$ g_1 = \frac{1}{2 + \frac{1}{ 3 + \frac{1}{4}}}$$

Here $n=2$,  $a_0 = 2, a_1 = 3, a_2 = 4$. This represents the number $0.43333\cdots$, which is the fraction $\frac{13}{30}$.

## Continued Fraction Grammar.

We will define continued fractions inductively by the grammar: 

$$\newcommand\nt[1]{\mathbf{#1}}
\newcommand\term[1]{\mathit{#1}}
$$

$$\begin{array}{rcl}
\nt{CFrac} & \Rightarrow & \term{OneByNumber}(\nt{Integer}) \\
 & | & \term{NumPlusContFrac}(\nt{Integer}, \nt{CFrac}) \\
 \end{array}$$
 
 Here $\nt{Integer}$ represents the scala integer number type (positive/negative whole numbers).
 
#### Example 1

The number $g_1$ above: 

$$ g_1 = \frac{1}{2 + \frac{1}{ 3 + \frac{1}{4}}}$$

will be represented in our grammar by our term: 

$$\term{NumPlusContFrac}(2, \term{NumPlusContFrac}(3,  \term{OneByNumber}(4))).$$

#### Example 2

The number $ g_2 = \frac{1}{6}$ will be represented in our grammar by our term:  $ \term{OneByNumber}(6).$


### Problem 1A (10 points)

Write down an inductive definition of the grammar above in Scala. Use the same names for terminals/nonterminals as in the grammar. Map the non terminal $\nt{Integer}$ to the built-in type $\mathsf{Int}$ in Scala.


<!-- grade_id: cell-d1659bc5ad17ffb6, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
sealed trait Cfrac
case class OneByNumber(divisor: Int) extends Cfrac
case class NumPlusContFrac(intgr: Int, frac: Cfrac) extends Cfrac
```

Output:
```
defined [32mtrait[39m [36mCfrac[39m
defined [32mclass[39m [36mOneByNumber[39m
defined [32mclass[39m [36mNumPlusContFrac[39m
```

<!-- grade_id: cell-ff1166c66661cf18, points: 10, grade: true, solution: false -->
```scala
//BEGIN TESTS
val g1 = NumPlusContFrac(2,NumPlusContFrac(3,NumPlusContFrac(4,OneByNumber(5))))
val g2 = OneByNumber(2)
val g3 = NumPlusContFrac(3, NumPlusContFrac(-1, OneByNumber(4)))
val g4 = NumPlusContFrac(1, OneByNumber(-2))
passed(10)
//END TESTS
```

Output:
```
*** Tests Passed (10 points) ***
[36mg1[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m(
  [32m2[39m,
  [33mNumPlusContFrac[39m([32m3[39m, [33mNumPlusContFrac[39m([32m4[39m, [33mOneByNumber[39m([32m5[39m)))
)
[36mg2[39m: [32mOneByNumber[39m = [33mOneByNumber[39m([32m2[39m)
[36mg3[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m3[39m, [33mNumPlusContFrac[39m([32m-1[39m, [33mOneByNumber[39m([32m4[39m)))
[36mg4[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m1[39m, [33mOneByNumber[39m([32m-2[39m))
```

### Problem 1B (5 points)
Implement a function `double_of_cfrac` that takes in a continued fraction expressed using the inductive definition above and evaluates the corresponding double precision number. 

#### Restrictions 

  - Please use case pattern matching. 
  - Your function _need not_ be tail recursive.
  - Forbidden: use of loops, mutable (vars), return statement.



<!-- grade_id: cell-1f8035faa2440711, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def double_of_cfrac(cont_frac: Cfrac): Double = cont_frac match{
    //base case return 1/divisor convert to double 
    case OneByNumber(divisor) => 1/divisor.toDouble 
    //else continue fraction, convert to double
    case NumPlusContFrac(num, constfrac) => {
        1 / (num.toDouble + double_of_cfrac(constfrac) )
    }
    
    //incase we are missing a case lol
    case _ => throw new IllegalArgumentException("Unhandled Case.")
}
```

Output:
```
defined [32mfunction[39m [36mdouble_of_cfrac[39m
```

<!-- grade_id: cell-8e8c3a4afc6890ae, points: 5, grade: true, solution: false -->
```scala
val g1 = NumPlusContFrac(2,NumPlusContFrac(3,OneByNumber(4)))
testWithMessage(double_of_cfrac(g1), 0.43333333, " # 1")
val g2 = OneByNumber(2)
testWithMessage(double_of_cfrac(g2), 0.5, " # 2")
val g3 = NumPlusContFrac(3, NumPlusContFrac(-1, OneByNumber(5)))
testWithMessage(double_of_cfrac(g3), 0.57142857, " # 3")
val g4 = NumPlusContFrac(1, OneByNumber(-2))
testWithMessage(double_of_cfrac(g4), 2.0, " # 4")

passed(10)
```

Output:
```
Test  # 1 -- comparing with tolerance 1.0E-5.
	 Expected: 0.43333333, your code returned: 0.43333333333333335
	 Passed!
Test  # 2 -- comparing with tolerance 1.0E-5.
	 Expected: 0.5, your code returned: 0.5
	 Passed!
Test  # 3 -- comparing with tolerance 1.0E-5.
	 Expected: 0.57142857, your code returned: 0.5714285714285714
	 Passed!
Test  # 4 -- comparing with tolerance 1.0E-5.
	 Expected: 2.0, your code returned: 2.0
	 Passed!

*** Tests Passed (10 points) ***
[36mg1[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m2[39m, [33mNumPlusContFrac[39m([32m3[39m, [33mOneByNumber[39m([32m4[39m)))
[36mg2[39m: [32mOneByNumber[39m = [33mOneByNumber[39m([32m2[39m)
[36mg3[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m3[39m, [33mNumPlusContFrac[39m([32m-1[39m, [33mOneByNumber[39m([32m5[39m)))
[36mg4[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m1[39m, [33mOneByNumber[39m([32m-2[39m))
```

### Problem 1C (10 points)

Write a function `ratio_of_cfrac` that takes in a continued fraction as defined by the grammar and returns a ratio $\frac{n}{d}$ where $n$ and $d$ are two integers such that the value of the continued fraction equals $\frac{n}{d}$. It is not required to reduce them in the lowest terms. The return type of your function must be a tuple of type `(Int, Int)`.

#### Restrictions

  - Please use case pattern matching.
  - Your function need not be tail recursive.
  - Forbidden: use of loops, mutable (vars), return statement.

<!-- grade_id: cell-6813d83b6b2177c7, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def ratio_of_cfrac(cont_frac: Cfrac) : (Int, Int) = cont_frac match{
    case OneByNumber(divisor) => (1, divisor)
    case NumPlusContFrac(num, contfrac) => {
        //get values for n and d
        val (n,d) = ratio_of_cfrac(contfrac)
        //was ((num*d + n), d) , forgot should return recipricol
        (d,(num*d + n)) 
    }
    //incase we are missing a case lol
    case _ => throw new IllegalArgumentException("Unhandled Case.")
    
}
```

Output:
```
defined [32mfunction[39m [36mratio_of_cfrac[39m
```

<!-- grade_id: cell-8687780755334866, points: 2, grade: true, solution: false -->
```scala
val g0 = OneByNumber(5)
val (n, d) = ratio_of_cfrac(g0)
println(s"Your Code Returned: ${n}/$d, Expected: 1/5")
assert(5 * n == d, "Test # 0 Failed.")
passed(2)
```

Output:
```
Your Code Returned: 1/5, Expected: 1/5

*** Tests Passed (2 points) ***
[36mg0[39m: [32mOneByNumber[39m = [33mOneByNumber[39m([32m5[39m)
[36mn[39m: [32mInt[39m = [32m1[39m
[36md[39m: [32mInt[39m = [32m5[39m
```

<!-- grade_id: cell-05e2a477771f6033, points: 2, grade: true, solution: false -->
```scala
val g1 = NumPlusContFrac(2,NumPlusContFrac(3,OneByNumber(4)))
val (n,d)= ratio_of_cfrac(g1)
//n/d == 13/30 ==> 30 * n == 13 * d
println(s"Your code returned: ${n}/${d}. Expected : 13/30")
assert (30 * n == 13* d, "Test #1  failed.")
passed(2)
```

Output:
```
Your code returned: 13/30. Expected : 13/30

*** Tests Passed (2 points) ***
[36mg1[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m2[39m, [33mNumPlusContFrac[39m([32m3[39m, [33mOneByNumber[39m([32m4[39m)))
[36mn[39m: [32mInt[39m = [32m13[39m
[36md[39m: [32mInt[39m = [32m30[39m
```

<!-- grade_id: cell-2c3d402c75fc4fac, points: 3, grade: true, solution: false -->
```scala
val g2 = NumPlusContFrac(2,NumPlusContFrac(3,NumPlusContFrac(4, OneByNumber(5))))
val (n,d)= ratio_of_cfrac(g2)
//n/d == 68/157 
println(s"Your code returned: ${n}/${d}. Expected : 68/157")
assert (157 * n == 68* d, "Test #2  failed.")
passed(3)
```

Output:
```
Your code returned: 68/157. Expected : 68/157

*** Tests Passed (3 points) ***
[36mg2[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m(
  [32m2[39m,
  [33mNumPlusContFrac[39m([32m3[39m, [33mNumPlusContFrac[39m([32m4[39m, [33mOneByNumber[39m([32m5[39m)))
)
[36mn[39m: [32mInt[39m = [32m68[39m
[36md[39m: [32mInt[39m = [32m157[39m
```

<!-- grade_id: cell-7ebcf06365df79c1, points: 3, grade: true, solution: false -->
```scala
val g3 = NumPlusContFrac(1,
            NumPlusContFrac(1, 
               NumPlusContFrac(1, 
                 NumPlusContFrac(1, 
                    NumPlusContFrac(1,OneByNumber(1))))))
val (n,d)= ratio_of_cfrac(g3)
//n/d == 8/13
println(s"Your code returned: ${n}/${d}. Expected : 8/13")
assert(13 * n == 8 * d, "Test # 3 Failed")
passed(3)
```

Output:
```
Your code returned: 8/13. Expected : 8/13

*** Tests Passed (3 points) ***
[36mg3[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m(
  [32m1[39m,
  [33mNumPlusContFrac[39m(
    [32m1[39m,
    [33mNumPlusContFrac[39m([32m1[39m, [33mNumPlusContFrac[39m([32m1[39m, [33mNumPlusContFrac[39m([32m1[39m, [33mOneByNumber[39m([32m1[39m))))
  )
)
[36mn[39m: [32mInt[39m = [32m8[39m
[36md[39m: [32mInt[39m = [32m13[39m
```

### Problem 1 D (10 points)

Write a function `cfrac_of_list` that given a list of integers `lst` of the form `List(a0, a1, ..., an)`, converts it into the continued fraction of depth $n$ given by the numbers `a0, .., an` in the list.

Function takes in a single argument whose 
 type must be `List[Int]`.
Return type of your function must be `CFrac`. 

If your input is empty your code must throw an `IllegalArgumentException`.


#### Example 1

Input: `List(5)`
Function returns: `OneByNumber(5)`

#### Example 2
Input: `List(2, 3, 4)`
Function returns: `NumPlusContFrac(2,NumPlusContFrac(3,OneByNumber(4)))`

#### Restrictions
 - No loops, var or return.
 - You can use pattern matching over Scala Lists and any list API functions of your choice.
 - Your code must throw an `IllegalArgumentException` if the list is empty.
 

<!-- grade_id: cell-f5d78dcab32f5687, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def cfrac_of_list(lst: List[Int]): Cfrac = lst match{ 
    //use guards to get rid of if else statements
    
    //empty list case
    case lst if(lst.isEmpty == true)  => throw new IllegalArgumentException("list cannot be empty")
    // list of one element 
    case list if(lst.length == 1) =>  OneByNumber(lst.head)
    // rest of list 
    case list if(lst.length > 1) => NumPlusContFrac(lst.head, cfrac_of_list(lst.tail))

}
```

Output:
```
defined [32mfunction[39m [36mcfrac_of_list[39m
```

<!-- grade_id: cell-7499fb1e6729f7d6, points: 2, grade: true, solution: false -->
```scala
val g0 = OneByNumber(5)
testWithMessage(cfrac_of_list(List(5)), g0, "Test # 1")
passed(2)
```

Output:
```
Test Test # 1
	 Expected: OneByNumber(5), your code returned: OneByNumber(5)
	 Passed!

*** Tests Passed (2 points) ***
[36mg0[39m: [32mOneByNumber[39m = [33mOneByNumber[39m([32m5[39m)
```

<!-- grade_id: cell-1a80f2fb7454fe99, points: 3, grade: true, solution: false -->
```scala
val g1 = NumPlusContFrac(1,NumPlusContFrac(2,OneByNumber(3)))
testWithMessage(cfrac_of_list(List(1,2,3)), g1, "Test # 2")
passed(3)
```

Output:
```
Test Test # 2
	 Expected: NumPlusContFrac(1,NumPlusContFrac(2,OneByNumber(3))), your code returned: NumPlusContFrac(1,NumPlusContFrac(2,OneByNumber(3)))
	 Passed!

*** Tests Passed (3 points) ***
[36mg1[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m([32m1[39m, [33mNumPlusContFrac[39m([32m2[39m, [33mOneByNumber[39m([32m3[39m)))
```

<!-- grade_id: cell-38afdb2eb86ce2a4, points: 2, grade: true, solution: false -->
```scala
try {
    println("Calling your function cfrac_of_list on empty list.")
    val g = cfrac_of_list(Nil);
    assert(false, "No IllegalArgumentException thrown on empty list")
} catch {
  case e:IllegalArgumentException => {
     println("IllegalArgumentException caught: Expected behavior seen!")
 }
  case e: Exception => {println(s"Wrong type of exception thrown-- $e")}
}

passed(3)
```

Output:
```
Calling your function cfrac_of_list on empty list.
IllegalArgumentException caught: Expected behavior seen!

*** Tests Passed (3 points) ***
```

<!-- grade_id: cell-1c6f18c32b93c31a, points: 3, grade: true, solution: false -->
```scala
val g3 = NumPlusContFrac(1,
            NumPlusContFrac(1, 
               NumPlusContFrac(1, 
                 NumPlusContFrac(1, 
                    NumPlusContFrac(1,OneByNumber(1))))))
testWithMessage(cfrac_of_list(List(1,1,1,1,1,1)), g3, "Test # 3")
passed(3)
```

Output:
```
Test Test # 3
	 Expected: NumPlusContFrac(1,NumPlusContFrac(1,NumPlusContFrac(1,NumPlusContFrac(1,NumPlusContFrac(1,OneByNumber(1)))))), your code returned: NumPlusContFrac(1,NumPlusContFrac(1,NumPlusContFrac(1,NumPlusContFrac(1,NumPlusContFrac(1,OneByNumber(1))))))
	 Passed!

*** Tests Passed (3 points) ***
[36mg3[39m: [32mNumPlusContFrac[39m = [33mNumPlusContFrac[39m(
  [32m1[39m,
  [33mNumPlusContFrac[39m(
    [32m1[39m,
    [33mNumPlusContFrac[39m([32m1[39m, [33mNumPlusContFrac[39m([32m1[39m, [33mNumPlusContFrac[39m([32m1[39m, [33mOneByNumber[39m([32m1[39m))))
  )
)
```

## Problem 2

In class, we introduced you to the inductive definition of __NumList__ (lists with numbers in them).

$$\begin{array}{rcl}
\nt{NumList} & \Rightarrow & \term{EmptyList} \\
& | & \term{ConsElt}(\nt{Integer}, \term{NumList}) \\
\end{array}$$

Suppose we extend the grammar by adding two additional rules to the ones already shown above.
$$\begin{array}{rcl}
\nt{NumList} & \Rightarrow & \term{ListConcat}(\nt{NumList}, \nt{NumList}) \\
& | & \term{AppendElt}(\nt{NumList}, \nt{Integer}) \\
\end{array}$$
Here _ListConcat_ models the concatenation of two __NumList__ and
_AppendElt_ models the appending of an element to the end of a __NumList__ (as opposed to _ConsElt_ that appends to the beginning).


### 2 A (5 points) 
Implement the Scala inductive definition of the type __NumList__ following the four grammar rules provided above. Ensure that the non-terminal __Integer__ is mapped to the Scala type $\mathsf{Int}$.

<!-- grade_id: cell-8b2a3993c4623861, points: 0, grade: false, solution: true -->
```scala
sealed trait NumList
case object EmptyList extends NumList
case class ConsElt(elt: Int, lst: NumList) extends NumList
case class ListConcat(lst1: NumList, lst2: NumList) extends NumList
case class AppendElt(lst1: NumList, appender: Int) extends NumList
```

Output:
```
defined [32mtrait[39m [36mNumList[39m
defined [32mobject[39m [36mEmptyList[39m
defined [32mclass[39m [36mConsElt[39m
defined [32mclass[39m [36mListConcat[39m
defined [32mclass[39m [36mAppendElt[39m
```

<!-- grade_id: cell-10401c115fcd5287, points: 5, grade: true, solution: false -->
```scala
val l1 = EmptyList 
val l2 = ConsElt(2, ConsElt(3, l1))
val l3 = ConsElt(1, ConsElt(2, l2))
val l4 = ListConcat(l2, l3)
val l5 = ConsElt(3, ListConcat(l2, l3))
val l6 = ConsElt(5, AppendElt(l5, 4))
val l7 = ConsElt(6, ListConcat(l6, EmptyList))
passed(5)
```

Output:
```
*** Tests Passed (5 points) ***
[36ml1[39m: [32mEmptyList[39m.type = EmptyList
[36ml2[39m: [32mConsElt[39m = [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList))
[36ml3[39m: [32mConsElt[39m = [33mConsElt[39m([32m1[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList))))
[36ml4[39m: [32mListConcat[39m = [33mListConcat[39m(
  [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList)),
  [33mConsElt[39m([32m1[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList))))
)
[36ml5[39m: [32mConsElt[39m = [33mConsElt[39m(
  [32m3[39m,
  [33mListConcat[39m(
    [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList)),
    [33mConsElt[39m([32m1[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList))))
  )
)
[36ml6[39m: [32mConsElt[39m = [33mC
```

### 2 B (10 points) 

Write a function `scala_list_of_numlist` that given an input of type `NumList` that you defined above, returns a corresponding scala list of type `List[Int]`.

#### Example

Input: 

~~~
 ListConcat(
  ConsElt(2, ConsElt(3, EmptyList)),
  ConsElt(1, ConsElt(2, ConsElt(2, ConsElt(3, EmptyList))))
)
~~~

Output: `List(2, 3, 1, 2, 2, 3)`

#### Restrictions

- No loops, var or return.
- List operations `::`, `++` and `:+` are allowed, and use of list API functions are allowed as well.


<!-- grade_id: cell-d62cd672a331c23f, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def scala_list_of_numlist(lst: NumList): List[Int] = lst match{
    case EmptyList => Nil
    case ConsElt(elt, lst) => elt :: scala_list_of_numlist(lst)
    case ListConcat(numlst1, numlst2) => scala_list_of_numlist(numlst1) ::: scala_list_of_numlist(numlst2)
    case AppendElt(numlst, appender) => scala_list_of_numlist(numlst) :+ appender
}
```

Output:
```
defined [32mfunction[39m [36mscala_list_of_numlist[39m
```

<!-- grade_id: cell-34bcbbcc3ce0addc, points: 3, grade: true, solution: false -->
```scala
val l1 = ListConcat(
  ConsElt(2, ConsElt(3, EmptyList)),
  ConsElt(1, ConsElt(2, ConsElt(2, ConsElt(3, EmptyList))))
)

testWithMessage(scala_list_of_numlist(l1), List(2, 3, 1, 2, 2, 3), "#1")
passed(3)
```

Output:
```
Test #1
	 Expected: List(2, 3, 1, 2, 2, 3), your code returned: List(2, 3, 1, 2, 2, 3)
	 Passed!

*** Tests Passed (3 points) ***
[36ml1[39m: [32mListConcat[39m = [33mListConcat[39m(
  [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList)),
  [33mConsElt[39m([32m1[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList))))
)
```

<!-- grade_id: cell-e91b81fa277e2d22, points: 1, grade: true, solution: false -->
```scala
val l2 = EmptyList
testWithMessage(scala_list_of_numlist(l2), Nil, "#2")
passed(1)
```

Output:
```
Test #2
	 Expected: List(), your code returned: List()
	 Passed!

*** Tests Passed (1 point) ***
[36ml2[39m: [32mEmptyList[39m.type = EmptyList
```

<!-- grade_id: cell-e7c4b79aa26999eb, points: 3, grade: true, solution: false -->
```scala
val l6 = ConsElt(
  5,
  AppendElt(
    ConsElt(
      3,
      ListConcat(
        ConsElt(2, ConsElt(3, EmptyList)),
        ConsElt(1, ConsElt(2, ConsElt(2, ConsElt(3, EmptyList))))
      )
    ),
    4
  )
)
testWithMessage(scala_list_of_numlist(l6), List(5, 3, 2, 3, 1, 2, 2, 3, 4), "#2")
passed(3)
```

Output:
```
Test #2
	 Expected: List(5, 3, 2, 3, 1, 2, 2, 3, 4), your code returned: List(5, 3, 2, 3, 1, 2, 2, 3, 4)
	 Passed!

*** Tests Passed (3 points) ***
[36ml6[39m: [32mConsElt[39m = [33mConsElt[39m(
  [32m5[39m,
  [33mAppendElt[39m(
    [33mConsElt[39m(
      [32m3[39m,
      [33mListConcat[39m(
        [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList)),
        [33mConsElt[39m([32m1[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList))))
      )
    ),
    [32m4[39m
  )
)
```

<!-- grade_id: cell-396db004ea367ee7, points: 3, grade: true, solution: false -->
```scala
val l7 =  ConsElt(
  6,
  ListConcat(
    ConsElt(
      5,
      AppendElt(
        ConsElt(
          3,
          ListConcat(
            ConsElt(2, ConsElt(3, EmptyList)),
            ConsElt(1, ConsElt(2, ConsElt(3, ConsElt(-1, EmptyList))))
          )
        ),
        -4
      )
    ),
    EmptyList
  )
)

testWithMessage(scala_list_of_numlist(l7), List(6, 5, 3, 2, 3, 1, 2, 3, -1, -4), "#2")
passed(3)
```

Output:
```
Test #2
	 Expected: List(6, 5, 3, 2, 3, 1, 2, 3, -1, -4), your code returned: List(6, 5, 3, 2, 3, 1, 2, 3, -1, -4)
	 Passed!

*** Tests Passed (3 points) ***
[36ml7[39m: [32mConsElt[39m = [33mConsElt[39m(
  [32m6[39m,
  [33mListConcat[39m(
    [33mConsElt[39m(
      [32m5[39m,
      [33mAppendElt[39m(
        [33mConsElt[39m(
          [32m3[39m,
          [33mListConcat[39m(
            [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, EmptyList)),
            [33mConsElt[39m([32m1[39m, [33mConsElt[39m([32m2[39m, [33mConsElt[39m([32m3[39m, [33mConsElt[39m([32m-1[39m, EmptyList))))
          )
        ),
        [32m-4[39m
      )
    ),
    EmptyList
  )
)
```

## That's All Folks!